# **Важно!**

Домашнее задание состоит из нескольких задач, которые вам нужно решить.
*   Баллы выставляются по принципу выполнено/невыполнено.
*   За каждую выполненую задачу вы получаете баллы (количество баллов за задание указано в скобках).

**Инструкция:** выполните задания в этом же ноутбуке (места под решения **каждой** задачи обозначаются как **#НАЧАЛО ВАШЕГО РЕШЕНИЯ** и **#КОНЕЦ ВАШЕГО РЕШЕНИЯ**).

**Как отправить задание на проверку:** вам необходимо сохранить решение в данном блокноте и отправить итоговый **файл .IPYNB** на учебной платформе в **стандартную форму сдачи домашнего задания.**

**Сроки проверки:** преподаватель проверит домашнее задание в течение недели после дедлайна и даст вам обратную связь.

# **Перед выполнением задания**

1. **Перезапустите ядро (restart the kernel):** в меню выбрать **Ядро (Kernel)**
→ **Перезапустить (Restart).**
2. **Выполните** **все ячейки (run all cells)**: в меню выбрать **Ячейка (Cell)**
→ **Запустить все (Run All)**.

После ячеек с заданием следуют ячейки с проверкой **с помощью assert**: если в коде есть ошибки, assert выведет уведомление об ошибке; если в коде нет ошибок, assert отработает без вывода дополнительной информации.

---

# Домашнее задание 6. Бустинг. XGBoost, CatBoost. Оптимизация гиперпараметров. Grid Search, Hyperopt

**Цели:**
*   Применить на практике такие методы, как XGBoost и CatBoost.
*   Оптимизировать гиперпараметры простым перебором (Grid Search).
*   Оптимизировать гиперпараметры с помощью пакета Hyperopt.

# Условия домашней работы

## Часть 1. XGBoost и Grid Search

**XGBoost (eXtreme Gradient Boosting)** — это оптимизированная реализация алгоритма градиентного бустинга над решающими деревьями. Он является одним из наиболее популярных и эффективных алгоритмов машинного обучения, широко применяемым для задач классификации и регрессии.

XGBoost использует ансамбль деревьев решений, которые обучаются последовательно, каждое следующее дерево исправляет ошибки предыдущих. Основная идея заключается в том, чтобы строить деревья, которые учитывают остатки (разницу между предсказанными и фактическими значениями) предыдущих деревьев. Это позволяет модели постепенно улучшать свою способность к прогнозированию и достигать лучшей обобщающей способности.

XGBoost предоставляет множество параметров, которые можно настраивать для достижения оптимальной производительности и обобщающей способности модели. Он обладает высокой скоростью обучения и предсказания, хорошей обработкой категориальных признаков, встроенной обработкой пропущенных значений и возможностью регуляризации для снижения переобучения.

Установка XGBoost:
`pip install xgboost`

Для обработки датасета в XGBoost требуется выполнить следующие преобразования:

1. Разделите данные на признаки (X) и целевую переменную (y). Признаки должны быть представлены в виде двумерного массива (например, в виде списка списков или в формате pandas DataFrame), а целевая переменная — в виде одномерного массива или списка.

2. Преобразуйте категориальные признаки в числовой формат. XGBoost не может обрабатывать категориальные данные напрямую, поэтому необходимо выполнить кодирование категориальных признаков. Вы можете использовать методы, такие как One-Hot Encoding или Label Encoding, чтобы преобразовать категориальные значения в числовые.

   - One-Hot Encoding — создает дополнительные бинарные признаки для каждой категории и присваивает им значение 0 или 1, в зависимости от принадлежности исходного значения к соответствующей категории. Многие библиотеки, включая pandas и scikit-learn, предоставляют инструменты для выполнения One-Hot Encoding.
   
   - Label Encoding — присваивает каждой уникальной категории целочисленное значение. Например, «красный» может быть закодирован как 0, «синий» — как 1 и т. д. В библиотеке scikit-learn есть класс LabelEncoder, который можно использовать для выполнения Label Encoding.

**GridSearchCV** — это инструмент в библиотеке scikit-learn, который предоставляет функциональность для систематического поиска наилучших гиперпараметров модели машинного обучения. Гиперпараметры являются настройками модели, которые не могут быть обучены в процессе обучения, а должны быть заданы до начала обучения. Примерами гиперпараметров могут быть глубина дерева в решающем дереве, скорость обучения в алгоритмах градиентного спуска и степень регуляризации в линейных моделях.

GridSearchCV позволяет определить набор возможных значений гиперпараметров, которые вы хотите проверить, и автоматически перебирает все возможные комбинации этих значений, обучая и оценивая модель на каждой комбинации. Затем он возвращает модель с наилучшими найденными значениями гиперпараметров.

В процессе использования GridSearchCV вам необходимо предоставить следующие элементы:
* Модель машинного обучения, которую вы хотите настроить с помощью поиска по сетке гиперпараметров.
* Словарь или список возможных значений для каждого гиперпараметра, которые вы хотите проверить.
* Метод оценки, который будет использоваться для сравнения различных комбинаций гиперпараметров. Например, это может быть точность (Accuracy) или площадь под ROC-кривой (AUC-ROC).

GridSearchCV автоматически перебирает все возможные комбинации гиперпараметров, выполняет обучение и оценку модели на каждой комбинации с использованием выбранного метода оценки. Затем он возвращает наилучшую модель с оптимальными значениями гиперпараметров, которые были найдены в результате перебора. Эта модель может быть использована для предсказания на новых данных.

```python
# Импорт необходимых библиотек
from sklearn.preprocessing import LabelEncoder
from sklearn.model_selection import train_test_split, GridSearchCV
from sklearn.metrics import accuracy_score
import xgboost as xgb
from xgboost import XGBClassifier
import pandas as pd
import numpy as np
```

```python
data = pd.read_csv("data/Telco-Customer-Churn.csv", sep=',')
```

```python
X = data.drop(['customerID', 'Churn'], axis=1)
y = data['Churn']
# Целевая переменная для модели XGBoost должна быть представлена целым числом.
y = data['Churn'].apply(lambda x: 1 if x == "Yes" else 0)
```

**Задание 1 (2 балла)**

1. Выполните для столбца TotalCharges преобразование к вещественному типу.

    Пример
    df["column_name"] = pd.to_numeric(X["column_name"], errors="coerce")

2. Создайте пустой список categorical_columns

3. Создайте цикл для прохода по каждому столбцу в DataFrame X:
    если количество уникальных элементов в столбце меньше 5:
        добавить имя столбца в categorical_columns.

    Пример:    
    for column in df.columns:
        if len(df[column].unique()) < 10:
            list_to_add.append(column)  

4. Создайте экземпляр класса LabelEncoder, сохраните его в label_encoder.

    Пример:
    encoder = LabelEncoder()

5. Примените LabelEncoder к каждому категориальному признаку.

    Пример:
    for column in categorical_columns:
        # Замена значений в столбце X[column] на их числовое представление:
        df[column] = label_encoder.fit_transform(df[column])

<!-- grade_id: cell-c93c03a36b18e7a1, points: 0, grade: false, solution: true -->
```python
# НАЧАЛО ВАШЕГО РЕШЕНИЯ
# 1. Преобразуем TotalCharges к вещественному типу
X["TotalCharges"] = pd.to_numeric(X["TotalCharges"], errors="coerce")

# 2. Создаём пустой список для категориальных колонок
categorical_columns = []

# 3. Проходим по всем столбцам и добавляем те, у которых < 5 уникальных значений
for column in X.columns:
    if len(data[column].unique()) < 5:
        categorical_columns.append(column)

# 4. Создаём LabelEncoder
from sklearn.preprocessing import LabelEncoder
label_encoder = LabelEncoder()

# 5. Применяем LabelEncoder ко всем категориальным признакам
for column in categorical_columns:
    X[column] = label_encoder.fit_transform(X[column])  # преобразуем в строку на случай NaN

# Покажем результат
X[categorical_columns].head()
# КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

Output:
```
gender  SeniorCitizen  Partner  Dependents  PhoneService  MultipleLines  \
0       0              0        1           0             0              1   
1       1              0        0           0             1              0   
2       1              0        0           0             1              0   
3       1              0        0           0             0              1   
4       0              0        0           0             1              0   

   InternetService  OnlineSecurity  OnlineBackup  DeviceProtection  \
0                0               0             2                 0   
1                0               2             0                 2   
2                0               2             2                 0   
3                0               2             0                 2   
4                1               0             0                 0   

   TechSupport  StreamingTV  StreamingMovies  Contract  PaperlessBilling  \
0            0            0          
```

<!-- grade_id: cell-67896dd84ccce32c, points: 2, grade: true, solution: false -->
```python
assert X["TotalCharges"].dtype == float
assert len(categorical_columns) == 16
```

**Задание 2 (1 балл)**

Разбейте датасет X и метки y, сохранив результат в X_train, X_test, y_train, y_test
    с помощью train_test_split. Параметры train_test_split: X, y, test_size=0.2, random_state=42.

<!-- grade_id: cell-fb424e10a55c73b7, points: 0, grade: false, solution: true -->
```python
# НАЧАЛО ВАШЕГО РЕШЕНИЯ
# Разделение на обучающую и тестовую выборки
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.2, random_state=42)
# КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

<!-- grade_id: cell-305e029df98ef617, points: 1, grade: true, solution: false -->
```python
assert len(X_train) == len(y_train) == 5634
```

**Задание 3 (1 балл)**

Инициализируйте экземпляр класса XGBClassifier без параметров и сохраните
его в переменную `xgb_classifier`.

<!-- grade_id: cell-fcf2e11d7fcf7326, points: 0, grade: false, solution: true -->
```python
# НАЧАЛО ВАШЕГО РЕШЕНИЯ
xgb_classifier = XGBClassifier()
# КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

<!-- grade_id: cell-5071f27eb1747157, points: 1, grade: true, solution: false -->
```python
assert type(xgb_classifier) == xgb.XGBClassifier
```

**Задание 4 (1 балл)**

Определение сетки параметров для перебора.
Создайте словарь parameters со следующими парами ключ-значение:
max_depth - [3, 5, 7]
learning_rate - [0.1, 0.01, 0.001]
n_estimators - [30, 50, 100]

Пример:
params = {
    'key1': [1,2,3],
    'key2': [1,2,3],
    'key3': [1,2,3]
}

<!-- grade_id: cell-23246e1e99a69175, points: 0, grade: false, solution: true -->
```python
# НАЧАЛО ВАШЕГО РЕШЕНИЯ
parameters = {
    'max_depth': [3, 5, 7],
    'learning_rate': [0.1, 0.01, 0.001],
    'n_estimators': [30, 50, 100]
}
# КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

<!-- grade_id: cell-54081ac097d6ac1a, points: 1, grade: true, solution: false -->
```python
assert type(parameters) == dict
assert sorted(parameters['max_depth']) == [3, 5, 7]
assert sorted(parameters['learning_rate']) == [0.001, 0.01, 0.1]
assert sorted(parameters['n_estimators']) == [30, 50, 100]
```

**Задание 5 (1 балл)**

Создайте экземпляр класса GridSearchCV, передав в него ряд параметров:
estimator=xgb_classifier,
param_grid=parameters,
scoring='accuracy',
cv=5.

Сохраните в переменную grid_search.

Пример:
grid_search = GridSearchCV(par1=val1, par2=val2)

<!-- grade_id: cell-7582baae7cd33e40, points: 0, grade: false, solution: true -->
```python
# НАЧАЛО ВАШЕГО РЕШЕНИЯ
grid_search = GridSearchCV(
    estimator=xgb_classifier,
    param_grid=parameters,
    scoring='accuracy',
    cv=5
)
# КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

<!-- grade_id: cell-f36f8a554fc343cf, points: 1, grade: true, solution: false -->
```python
assert type(grid_search) == GridSearchCV
assert grid_search.__dict__['scoring'] == 'accuracy'
assert grid_search.__dict__['cv'] == 5
```

```python
# Обучение модели с перебором параметров
# !!! Может выполняться несколько минут
grid_search.fit(X_train, y_train)
```

Output:
```
GridSearchCV(cv=5,
             estimator=XGBClassifier(base_score=None, booster=None,
                                     callbacks=None, colsample_bylevel=None,
                                     colsample_bynode=None,
                                     colsample_bytree=None, device=None,
                                     early_stopping_rounds=None,
                                     enable_categorical=False, eval_metric=None,
                                     feature_types=None, feature_weights=None,
                                     gamma=None, grow_policy=None,
                                     importance_type=None,
                                     interaction_constraints=None,
                                     learning_rate=None, max_bin=None,
                                     max_cat_threshold=None,
                                     max_cat_to_onehot=None,
                                     max_delta_step=None, max_depth=None,
                  
```

```python
# Лучшие параметры и точность модели
best_params = grid_search.best_params_
best_score = grid_search.best_score_

print("Лучшие параметры:", best_params)
print("Точность на обучающем наборе:", best_score)

# Оценка модели на тестовом наборе
y_pred = grid_search.predict(X_test)
accuracy = accuracy_score(y_test, y_pred)
print("Точность на тестовом наборе:", accuracy)
```

Output:
```
Лучшие параметры: {'learning_rate': 0.1, 'max_depth': 3, 'n_estimators': 50}
Точность на обучающем наборе: 0.8026269462144267
Точность на тестовом наборе: 0.8133427963094393
```

## Часть 2. CatBoost и Hyperopt

**CatBoost** — это открытая библиотека градиентного бустинга на деревьях решений. Она разработана компанией Яндекс и предназначена для решения задач машинного обучения, таких как классификация, регрессия и ранжирование. CatBoost предлагает высокую скорость обучения и прогнозирования, а также хорошую устойчивость к различным типам данных и категориальным признакам.

Особенность CatBoost – способность автоматически обрабатывать категориальные признаки без необходимости их предварительного преобразования или использования кодирования, что делает эту библиотеку особенно полезной для работы с данными с множеством категориальных переменных.

Установка CatBoost:
`pip install catboost`

**Hyperopt** — это библиотека для автоматической оптимизации гиперпараметров моделей машинного обучения. Она предоставляет простой и гибкий интерфейс для поиска наилучших значений гиперпараметров, используя алгоритмы оптимизации.

Основная идея Hyperopt — применение адаптивной модели оптимизации для нахождения комбинации гиперпараметров, которая минимизирует или максимизирует заданную функцию оценки (например, функцию потерь или метрику качества). Библиотека предлагает несколько алгоритмов оптимизации, включая Tree of Parzen Estimators (TPE) и Random Search.

**Процесс использования Hyperopt:**

1. **Определение пространства поиска.** Задается пространство гиперпараметров, которое нужно оптимизировать. Это может быть сделано с использованием объектов из модуля `hyperopt.hp`, которые представляют различные типы параметров (например, вещественные числа, целые числа, категориальные значения и др.).

2. **Определение функции оценки (целевой функции).** Задается функция, которую нужно минимизировать или максимизировать. Обычно это функция потерь или метрика качества, связанная с моделью машинного обучения.

3. **Создание объекта `Trials`.** Создается объект `Trials`, который будет использоваться для хранения информации о проведенных оптимизациях.

4. **Запуск оптимизации.** Вызывается функция `fmin` с указанием функции оценки, пространства поиска и алгоритма оптимизации. Hyperopt будет итеративно предлагать наборы гиперпараметров, вычислять значение функции оценки и обновлять модель оптимизации для принятия решения о следующем наборе гиперпараметров.

5. **Получение лучших гиперпараметров.** По окончании оптимизации можно получить значения лучших гиперпараметров, найденных в ходе процесса оптимизации.

Установка библиотеки: `pip install hyperopt`

```python
!pip install catboost
```

```python
# Импортируем необходимые библиотеки
import catboost as cb
from catboost import CatBoostClassifier
from sklearn.model_selection import train_test_split
from sklearn.metrics import accuracy_score
from hyperopt import hp, fmin, tpe, Trials

# Загрузка данных
data = pd.read_csv('data/Telco-Customer-Churn.csv')
X = data.drop(['customerID', 'Churn'], axis=1)
y = data['Churn']

X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.2, random_state=42)
```

**Задание 6 (1 балл)**

1. Выполните для столбца TotalCharges преобразование к вещественному типу.

2. Создайте пустой список categorical_columns_indices.

3. Создайте цикл с enumerate для прохода по каждому столбцу в DataFrame X:
    если количество уникальных элементов в столбце меньше 5:
        добавить индекс столбца в categorical_columns_indices
    Пример:    
    for index, column in enumerate(df.columns):
        if len(df[column].unique()) < 10:
            list_to_add.append(index)

<!-- grade_id: cell-1bed802cedc0e8e7, points: 0, grade: false, solution: true -->
```python
# НАЧАЛО ВАШЕГО РЕШЕНИЯ
# 1. Преобразование TotalCharges
X["TotalCharges"] = pd.to_numeric(X["TotalCharges"], errors="coerce")

# 2. Пустой список для индексов категориальных колонок
categorical_columns_indices = []

# 3. Цикл по столбцам с использованием enumerate
for index, column in enumerate(X.columns):
    if X[column].nunique() < 5:
        categorical_columns_indices.append(index)
# КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

<!-- grade_id: cell-0a6d79b9bba4aa27, points: 1, grade: true, solution: false -->
```python
assert X["TotalCharges"].dtype == float
assert len(categorical_columns_indices) == 16
```

**Задание 7 (1 балл)**

Создайте словарь dict_of_params со следующими парами ключ — значение:

*   "iterations": [30, 50, 70]
*   "depth": [3, 5, 7]
*   "learning_rate": [0.01, 0.3]
*   "l2_leaf_reg": [1, 3, 5, 7, 9



<!-- grade_id: cell-025a19e639ff83bb, points: 0, grade: false, solution: true -->
```python
# НАЧАЛО ВАШЕГО РЕШЕНИЯ
dict_of_params = {
    "iterations": [30, 50, 70],
    "depth": [3, 5, 7],
    "learning_rate": [0.01, 0.3],
    "l2_leaf_reg": [1, 3, 5, 7, 9]
}
# КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

<!-- grade_id: cell-88eef32de85e06d3, points: 1, grade: true, solution: false -->
```python
assert sorted(dict_of_params["iterations"]) == [30, 50, 70]
```

```python
# Переводим словарь dict_of_params в формат Hyperopt
space = {
    'iterations': hp.choice('iterations', dict_of_params["iterations"]),
    'depth': hp.choice('depth', dict_of_params["depth"]),
    'learning_rate': hp.uniform('learning_rate', dict_of_params["learning_rate"][0], dict_of_params["learning_rate"][1]),
    'l2_leaf_reg': hp.choice('l2_leaf_reg', dict_of_params["l2_leaf_reg"])
}
```

**Задание 8 (2 балла)**

Определяем функцию для оптимизации с гиперпараметрами и параметром `cat_features`, указывающим индексы категориальных признаков.

1. Определите функцию с именем objective с параметром params.

    Пример :
    def objective(parameters):


2. **Внутри функции objective** инициализируйте `cb.CatBoostClassifier()`, передав в качестве параметров `**params`, `cat_features=categorical_columns_indices`. Сохраните экземпляр в переменную `model`.

    Пример:
    model = cb.CatBoostClassifier(**parameters, cat_features=indices)

3. Обучите модель на `X_train`, `y_train` с помощью метода `fit(X_train, y_train, silent=True)`. Флаг `silent=True` убирает вывод подробной информации об обучении.

4. Сделайте предсказания для `X_test` с помощью метода `predict()`, ответ сохраните в `y_pred`.

5. Сохраните в переменную `score` результат вызова `accuracy_score()` с параметрами `y_test`, `y_pred`.

6. Верните через `return -score`.

<!-- grade_id: cell-a6ff96f3b6fa8479, points: 0, grade: false, solution: true -->
```python
# НАЧАЛО ВАШЕГО РЕШЕНИЯ
def objective(params):
    model = CatBoostClassifier(
        **params,
        cat_features=categorical_columns_indices,
        silent=True
    )
    model.fit(X_train, y_train)

    y_pred = model.predict(X_test)

    score = accuracy_score(y_test, y_pred)
    return -score
# КОНЕЦ ВАШЕГО РЕШЕНИЯ
```

<!-- grade_id: cell-76a46666a33dada0, points: 2, grade: true, solution: false -->
```python
assert callable(objective)
```

```python
# Запускаем оптимизацию с помощью Hyperopt
trials = Trials()
best = fmin(objective, space, algo=tpe.suggest, max_evals=100, trials=trials)
```

Output:
```
100%|██████████| 100/100 [00:12<00:00,  8.18trial/s, best loss: -0.8246983676366217]
```

```python
# Получаем лучшие значения гиперпараметров
final_best_params = {}
for key in best:
    try:
        final_best_params[key] = dict_of_params[key][best[key]]
    except:
        final_best_params[key] = best[key]

# Обучаем модель с лучшими гиперпараметрами
# Создание экземпляра CatBoostClassifier с передачей списка индексов категориальных признаков
model = cb.CatBoostClassifier(**final_best_params, cat_features=categorical_columns_indices)
model.fit(X_train, y_train, silent=True)

# Предсказываем классы для тестовых данных
y_pred = model.predict(X_test)

# Вычисляем точность модели
accuracy = accuracy_score(y_test, y_pred)

# Выводим результаты
print("Точность модели:", accuracy)
```

Output:
```
Точность модели: 0.8246983676366217
```

Мы можем видеть, что подготовка данных для CatBoost занимает меньше времени.
Если сравнивать GridSearchCV и Hyperopt, то Hyperopt устроен сложнее, однако может выдавать более оптимальные результаты.

# Поздравляем!
В этом домашнем задании вы научились пользоваться библиотеками `XGBoost` и `CatBoost`, освоили автоматическую оптимизацию гиперпараметров с методом перебора (Grid Search), а также более элегантным образом с помощью пакета `Hyperopt`.